# Lab 1: Data Types, Functions and Conditionals (29 points)

#### Exercise 0
Execute the code and try to understand the error msg

```python
print('Hello World'
```

Output:
```
File "/tmp/ipykernel_1073/403058330.py", line 1
    print('Hello World'
                       ^
SyntaxError: incomplete input
```

```python
'Welcome'/2
```

```python
print(a)
```

#### Exercise 1
 Write a function which computes gross pay given the **hours** worked and the **rate_per_hour**.


<!-- grade_id: cell-312caeaaf6abe1a9, points: 0, grade: false, solution: true -->
```python
def get_gross_pay(hours, rate_per_hour):
  pay=hours*rate_per_hour
  return pay
assert get_gross_pay(3,10) == 30
```

<!-- grade_id: cell-458cb3b5b6333645, points: 1, grade: true, solution: false -->
```python
assert get_gross_pay(3, 10) == 30
```

#### Exercise 2
Assume that we execute the following assignment statements.
```
width = 17
height = 12.0
```
For each of the following expressions, write the value of the expression and the type (of the value of the expression).

1. `width//2`

2. `width/2.0`

3. `height/3`

4. `1 + 2 * 5`

Execute the above code to check.

```python
'''
1. = 8 int
2. = 8.5 float
3. = 4.0 float
4. = 11 int
'''
1 + 2*5
```

Output:
```
11
```

#### Exercise 3
Rewrite your `get_gross_pay` function from ch2, exercise 3 to give the employee 1.5 times the hourly rate for hours worked above 40 hours.

<!-- grade_id: cell-8f610a982f866fac, points: 0, grade: false, solution: true -->
```python
def get_gross_pay(hours, rate_per_hour):
    if hours > 40:
      pay= hours *1.5 * rate_per_hour
    else:
      pay = hours*rate_per_hour
    return pay
assert get_gross_pay(0, 1000)==0
assert get_gross_pay(2, 10)==20
```

<!-- grade_id: cell-537d9d1c5dcd290f, points: 1, grade: true, solution: false -->
```python
assert get_gross_pay(0, 1000)==0
assert get_gross_pay(2, 10)==20
```

#### Exercise 4
What is the purpose of the `def` keyword in Python?

a) It is slang that means “the following code is really cool”

b) It indicates the start of a function

c) It indicates that the following indented section of code is to be stored for later

d) b and c are both true

e) None of the above

<!-- grade_id: ch4-ex1-ans, points: 0, grade: false, solution: true -->
```python
# Enter your answer as a python string in a variable 'ans'.
# Should be one of ['a', 'b', 'c', 'd', 'e']
ans="d"
assert type(ans) == str
assert ans in ['a','b','c','d','e']
```

<!-- grade_id: ch4-ex1-test, points: 1, grade: true, solution: false -->
```python
assert type(ans) == str
assert ans in ['a','b','c','d','e']
```

#### Exercise 5
Create a new GitHub repository for the Data Science course and complete the following tasks:
### Set up the repository
Name the repository appropriately (e.g., "stats507-coursework");
Ensure the repository is public (unless otherwise specified)

### Add a README file
Create a file named "README.md" in the root directory;
Include a brief description of the course and the purpose of the repository

### Provide the repository link


<!-- grade_id: cell-9a484821ec48c9fc, points: 0, grade: false, solution: true -->
```python
https://github.com/stuzinja75/Data-Science-101/blob/main/README.md
raise NotImplementedError()
```